### Aufgabe 1

```python
import datetime
print(datetime.datetime.now())
```

<!-- grade_id: cell-28df1799f8f8b769, points: 0, grade: false, solution: true -->
```python
def reverse(s):
    ###BEGIN SOLUTION
    return s[::-1]
    ###END SOLUTION
```

```python
reverse('Test')
```

Output:
```
'tseT'
```

<!-- grade_id: cell-da8c82e850a1922b, points: 1, grade: true, solution: false -->
```python
assert (reverse('Test') == 'tseT')
```

<!-- grade_id: cell-81540a070d18c412, points: 1, grade: true, solution: false -->
```python
assert (reverse('lol') == 'lol')
```

## Aufgabe 2
What are "fake"-threads?

Answer:

## Aufgabe 3
Does Java use "fake"-threads? Explain why or why not?

<!-- grade_id: cell-9ea0264ada6c25bd, points: 2, grade: true, solution: true -->
