# Introduction to Python

## Variables

### Assignment NR 1 

Create five different variables

These should be as follows:

A: Integer

B: Float

C: String

D: Boolean

E: List

<!-- grade_id: cell-2b117bf24c2bc84f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
A= 2
B= 3.5
C= "Hello World"
D= False
E= ["A","B"]
### END SOLUTION
```

<!-- grade_id: cell-0d290c45e645f07e, points: 1, grade: true, solution: false -->
```python
assert isinstance(A, int)
```

<!-- grade_id: cell-b1fff53c641c7b0d, points: 1, grade: true, solution: false -->
```python
assert isinstance(B, float)
```

<!-- grade_id: cell-d4d59ed9672e2a86, points: 1, grade: true, solution: false -->
```python
assert isinstance(C, str)
```

<!-- grade_id: cell-520cb870ee432dfd, points: 1, grade: true, solution: false -->
```python
assert isinstance(D, bool)
```

<!-- grade_id: cell-46ad3dcda74281cc, points: 1, grade: true, solution: false -->
```python
assert isinstance(E, list)
```

## Simple Operations

### Assigment NR 2

Now we do some simple addition and subtraction 

Create two new variables

SUM: Add A and B together

MULTIPLY: Multiply A and B together

<!-- grade_id: cell-778548d1dd6f4342, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
SUM= A+B
MULTIPLY= A*B
### END SOLUTION
```

<!-- grade_id: cell-af1a2716bc031a75, points: 1, grade: true, solution: false -->
```python
assert (SUM == A+B)
```

<!-- grade_id: cell-5b84a812571001fb, points: 1, grade: true, solution: false -->
```python
assert (MULTIPLY == A*B)
```

### Assigment NR 3

One of the most important operations within Python is the print() statement

print out the result of your SUM and MULTIPLY Operation

Also print out the type of these operations using the type() function

<!-- grade_id: cell-75127fe96710e70a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
print (SUM)
print(MULTIPLY)
print(type(SUM))
print(type(MULTIPLY))
### END SOLUTION
```

Output:
```
5.5
7.0
<class 'float'>
<class 'float'>
```

<!-- grade_id: cell-d32a94a739815222, points: 2, grade: true, solution: false -->
```python
assert isinstance(SUM,float)
assert isinstance(MULTIPLY, float)
```

## Lists

### Assignment NR 4

Create a list PRIME with the following values and order: 2, 3, 5, 7, 11, 13, 17

Then the create a subset SUBTEN of all PRIME values smaller than 10

<!-- grade_id: cell-a2918293e2e9aaee, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
PRIME = [2,3,5,7,11,13,17]
SUBTEN = PRIME[0:4]
### END SOLUTION
```

<!-- grade_id: cell-6399decb4436256e, points: 2, grade: true, solution: false -->
```python
assert (PRIME == [2,3,5,7,11,13,17])
assert (SUBTEN == [2,3,5,7])
```

## Dictionaries

### Assignment NR 5

Create a dictionary TUWIEN that includes the following keys:

Name: TU Wien

Location: Vienna

Students: 42000

Founded: 1815



<!-- grade_id: cell-467c0c91da83c1c3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
TUWIEN = {'Name': 'TU Wien', 'Location': 'Vienna', 'Students': '42000', "Founded": 1815}
### END SOLUTION
```

<!-- grade_id: cell-d0629aa844a668eb, points: 0, grade: true, solution: false -->
```python
assert (TUWIEN=={'Name': 'TU Wien', 'Location': 'Vienna', 'Students': '42000', "Founded": 1815} )
```

### Assignment NR 6

Now let's update the number of students to 43000

Also delete the key "Founded"



<!-- grade_id: cell-1ae9c032144fab01, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
TUWIEN['Students']=43000
del TUWIEN['Founded']
### END SOLUTION
```

<!-- grade_id: cell-b59c617109cf7e1b, points: 0, grade: true, solution: false -->
```python
assert (TUWIEN['Students'] == 43000)
assert ('Founded' not in TUWIEN)
```

## Loops

### Assignment NR 7
In this assigment we will create both a for loop and a while loop

The Goal of the for loop will be to add all integers from 0 to 10

The result of the for loop should be saved within the variable SUM_FOR

The result of the while loop should be saved within the variables SUM_WHILE


<!-- grade_id: cell-b8943fa1ffa8bee7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
SUM_FOR = 0
for x in range(0,11):
    SUM_FOR += x


SUM_WHILE = 0
x=0
while x <= 10:
    SUM_WHILE = x + SUM_WHILE
    x = x + 1

### END SOLUTION
```

<!-- grade_id: cell-7ce8c4913f8d86a3, points: 0, grade: true, solution: false -->
```python
assert (SUM_FOR == 55)
assert (SUM_WHILE == 55)
```

## Conditions

### Assignment NR 8

Let's assume you want to check a list of strings for the lengths of the strings inside

All Strings shorter than or equal to 5 characters should be counted within the variable SHORT

All Strings longer than 5 characters will be counted within the variable LONG

TIP: you can find the length of an item using len(list[index])

Create a Loop that iterates through all items within any given list and check for the two given conditions.


<!-- grade_id: cell-8aa4f5e9e32f30e7, points: 0, grade: false, solution: true -->
```python
def List_Length_Check(list):
    SHORT = 0
    LONG = 0
    length = len(list)
    ### BEGIN SOLUTION
    for i in range(length):
        if len(list[i]) <=5:
            SHORT += 1
        else:
            LONG += 1
    ### END SOLUTION
    return SHORT, LONG
```

<!-- grade_id: cell-f57c122dd606ffba, points: 2, grade: true, solution: false -->
```python
assert (List_Length_Check(["adfadfdasf","adsf","adfadfdasf"]) == (1,2))

assert (List_Length_Check(["asd","adsf","asd"]) == (3,0))
```

## Functions

In this part you will create your own functions

### Assignment NR 9

create a function that takes in a list of measurements that were taken in m. Our goal is to convert the entire list into feet

TIP: conversion of meters to feet is 1 m = 3.28084 ft

<!-- grade_id: cell-8944638593b7eea6, points: 0, grade: false, solution: true -->
```python
def meter_to_feet(list):
    ### BEGIN SOLUTION
    x = len (list)
    for i in range(x):
        list[i] = list[i]*3.28084
    return list
    ### END SOLUTION
```

<!-- grade_id: cell-39d95876b5f20ec0, points: 2, grade: true, solution: false -->
```python
assert (meter_to_feet([1,5,6]) == [3.28084, 16.4042, 19.68504])
```

### Assignment NR 10

The final assignment for this capter will be to create a function that takes two arguments and create a list from these.

The first value will be an integer to which we will count.

The second will also be an integer.

The goal is to create a create list (starting from 1) of every integer within but omitting those that are divisible by the second one.

Example. First number 10, second number 5. Output should be [0,1,2,3,4,6,7,8,9]

<!-- grade_id: cell-6237173a2ffe17e2, points: 0, grade: true, solution: true -->
```python
def Division_By(first, second):
    new_list=[]
    ### BEGIN SOLUTION
    for i in range(first):
        if i%second == 0:
            continue
        else:
            new_list.append(i)
    return new_list
    ### END SOLUTION
```

<!-- grade_id: cell-6e4e44a8a35af91a, points: 0, grade: true, solution: false -->
```python
assert (Division_By(20,10)== [1, 2, 3, 4, 5, 6, 7, 8, 9, 11, 12, 13, 14, 15, 16, 17, 18, 19])
```